<a href="https://colab.research.google.com/github/sravanthi200719/HDS5210_InClass/blob/master/sravanthi_week05_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 5.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(strings):
    """(list) -> int
    This function finds the longest string in the list of input strings
    and returns the position number of that longest string.

    >>> longest(['happy', 'birthday', 'to', 'me'])
    1

    >>> longest(['one', 'two', 'three', 'four'])
    2
    """

    longest_pos = 0
    longest_len = 0

    print("Position  Item         Longest So Far")
    print("--------- ------------ --------------")
    for position, item in enumerate(strings):

        if len(item) > longest_len:
            longest_pos = position
            longest_len = len(item)
        print("{0:9} {1:12} {2} : {3}".format(position, item, longest_pos, longest_len))

    return longest_pos

### END SOLUTION
```

```python
def longest(strings):
    """(list) -> int
    This function finds the longest string in the list of input strings
    and returns the position number of that longest string.

    >>> longest(['happy', 'birthday', 'to', 'me'])
    1

    >>> longest(['one', 'two', 'three', 'four'])
    2

    >>> longest(['a', 'ab', 'abc', 'abcd'])
    3

    >>> longest(['apple', 'banana', 'cherry', 'date'])
    1
    """

    if not strings:
        return -1  # Return -1 for an empty list to indicate no position

    longest_pos = 0
    longest_len = len(strings[0])  # Initialize with the length of the first string

    for position, item in enumerate(strings):
        if len(item) > longest_len:
            longest_pos = position
            longest_len = len(item)

    return longest_pos

```

```python
def longest(my_list):
    """(list)->int
    Find the longest item in a list of strings and return the index for that string.
    If there are multiple strings with the same length, return the index of the first of those longest strings.

    >>> longest(['happy', 'birthday', 'to', 'me'])
    1

    >>> longest(['abc','def','ghi'])
    0
    """

    if type(my_list) != list:
        raise TypeError("That isn't a list!")

    if len(my_list) == 0:
        raise ValueError("There were no items in that list!")

    max_length = -1
    max_index = None

    for num, item in enumerate(my_list):
        try:
            if len(item) > max_length:
                max_length = len(item)
                max_index = num
        except:
            pass

    return max_index

```

```python
def longest(my_list):
    """(list) -> int
    Find the longest item in a list of strings and return the index for that string.
    If there are multiple strings with the same length, return the index of the first of those longest strings.

    >>> longest(['happy', 'birthday', 'to', 'me'])
    1

    >>> longest(['abc', 'def', 'ghi'])
    0

    >>> longest(['happy', 'birthday', 2, 'me'])
    1  # Ignores the integer and returns the index of the longest string.

    >>> longest([])
    -1  # No strings in the list.
    """

    if not isinstance(my_list, list):
        return "Input must be a list!"  # Return a message instead of raising an error

    max_length = -1
    max_index = -1

    for num, item in enumerate(my_list):
        if isinstance(item, str):  # Only consider string items
            if len(item) > max_length:
                max_length = len(item)
                max_index = num

    return max_index  # Will return -1 if no strings are found

```

```python
longest(['happy', 'birthday', 2, 'me'])
```

Output:
```
1
```

```python
longest(['happy', 'birthday', 'to', 'me'])
```

Output:
```
1
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
PYDEV DEBUGGER WARNING:
sys.settrace() should not be used when the debugger is being used.
This may cause the debugger to stop working correctly.
If this is needed, please check: 
http://pydev.blogspot.com/2007/06/why-cant-pydev-debugger-work-with.html
to see how to restore the debug tracing back correctly.
Call Location:
  File "/usr/lib/python3.10/doctest.py", line 1501, in run
    sys.settrace(save_trace)

Finding tests in NoName
Trying:
    longest(['happy', 'birthday', 'to', 'me'])
Expecting:
    1
ok
Trying:
    longest(['abc', 'def', 'ghi'])
Expecting:
    0
ok
Trying:
    longest(['happy', 'birthday', 2, 'me'])
Expecting:
    1  # Ignores the integer and returns the index of the longest string.
**********************************************************************
File "__main__", line 12, in NoName
Failed example:
    longest(['happy', 'birthday', 2, 'me'])
Expected:
    1  # Ignores the integer and returns the index of the longest string.
Got:
    1
Trying:
    longest([])
Ex
```

```python
### MY OWN TEST CASES  1.
def custom_tests():

    # Basic functionality test
    result = longest(['happy', 'birthday', 'to', 'me'])
    print("Test 1 - longest(['happy', 'birthday', 'to', 'me']):", result)  # Expected: 1

    # Test using split on a string
    result = longest("happy birthday to me".split(" "))
    print("Test 2 - longest('happy birthday to me'.split()):", result)  # Expected: 1

    result = longest("enjoy class".split(" "))
    print("Test 3 - longest('enjoy class'.split()):", result)  # Expected: 0

    # More string comparisons
    result = longest(['when', 'what', 'where', 'how', 'who'])
    print("Test 4 - longest(['when', 'what', 'where', 'how', 'who']):", result)  # Expected: 2

    # Additional test cases
    result = longest(['a', 'ab', 'abc', 'abcd'])
    print("Test 5 - longest(['a', 'ab', 'abc', 'abcd']):", result)  # Expected: 3

    result = longest(['apple', 'banana', 'kiwi', 'mango'])
    print("Test 6 - longest(['apple', 'banana', 'kiwi', 'mango']):", result)  # Expected: 1

    result = longest(['123', '4567', '89'])
    print("Test 7 - longest(['123', '4567', '89']):", result)  # Expected: 1

    # Edge case with empty strings
    result = longest(['', ''])
    print("Test 8 - longest(['', '']):", result)  # Expected: 0

    print("All custom tests executed!")

# Run custom tests
custom_tests()



```

Output:
```
Test 1 - longest(['happy', 'birthday', 'to', 'me']): 1
Test 2 - longest('happy birthday to me'.split()): 1
Test 3 - longest('enjoy class'.split()): 0
Test 4 - longest(['when', 'what', 'where', 'how', 'who']): 2
Test 5 - longest(['a', 'ab', 'abc', 'abcd']): 3
Test 6 - longest(['apple', 'banana', 'kiwi', 'mango']): 1
Test 7 - longest(['123', '4567', '89']): 1
Test 8 - longest(['', '']): 0
All custom tests executed!
```

```python
#  MY OWN TEST CASE for longest function 2.
result = longest(['elephant', 'cat', 'giraffe', 'hippopotamus'])
print("Test - longest(['elephant', 'cat', 'giraffe', 'hippopotamus']):", result)  # Expected: 3

```

Output:
```
Test - longest(['elephant', 'cat', 'giraffe', 'hippopotamus']): 3
```

### 5.2 Celsius to Farenheit

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION


### END SOLUTION
```

```python
def f_to_c(tempf):
    """(float) -> float
    Convert temperature from degrees Fahrenheit to degrees Celsius using the formula:

    tempc = (5/9) * (tempf - 32)

    >>> f_to_c(212)
    100.0

    >>> f_to_c(32)
    0.0

    >>> f_to_c(98.6)
    37.0
    """
    tempc = (5 / 9) * (tempf - 32)
    return tempc

# Running the tests
if __name__ == "__main__":
    import doctest
    doctest.testmod()

```

Output:
```
**********************************************************************
File "__main__", line 12, in __main__.longest
Failed example:
    longest(['happy', 'birthday', 2, 'me'])
Expected:
    1  # Ignores the integer and returns the index of the longest string.
Got:
    1
**********************************************************************
File "__main__", line 15, in __main__.longest
Failed example:
    longest([])
Expected:
    -1  # No strings in the list.
Got:
    -1
**********************************************************************
1 items had failures:
   2 of   4 in __main__.longest
***Test Failed*** 2 failures.
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
f_to_c(212)
```

Output:
```
100.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

```python
### MY OWN TEST CASES 1. .
def custom_tests():
    # Test cases
    test_cases = [
        (212, 100.0),    # Boiling point of water
        (32, 0.0),       # Freezing point of water
        (98.6, 37.0),    # Average body temperature
        (0, -17.7778),   # Below freezing
        (-40, -40.0),    # Same in both scales
        (100, 37.7778),  # Above freezing
        (68, 20.0),      # Room temperature
        (0, -17.7778)    # Freezing point
    ]

    for tempf, expected in test_cases:
        result = f_to_c(tempf)
        print(f"f_to_c({tempf}) = {result:.4f}, Expected: {expected:.4f}")

# Run custom tests
custom_tests()

```

Output:
```
f_to_c(212) = 100.0000, Expected: 100.0000
f_to_c(32) = 0.0000, Expected: 0.0000
f_to_c(98.6) = 37.0000, Expected: 37.0000
f_to_c(0) = -17.7778, Expected: -17.7778
f_to_c(-40) = -40.0000, Expected: -40.0000
f_to_c(100) = 37.7778, Expected: 37.7778
f_to_c(68) = 20.0000, Expected: 20.0000
f_to_c(0) = -17.7778, Expected: -17.7778
```

```python
## MY OWN TEWST CASES 2.
def additional_test_cases():
    # Additional test case
    tempf = 45
    expected = (5 / 9) * (tempf - 32)  # Calculating expected value
    result = f_to_c(tempf)
    print(f"f_to_c({tempf}) = {result:.4f}, Expected: {expected:.4f}")

# Run additional test case
additional_test_cases()

```

Output:
```
f_to_c(45) = 7.2222, Expected: 7.2222
```

### 5.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION

### END SOLUTION
```

```python
from datetime import date

def length_of_stay(encounters):
    """(dict) -> dict
    Calculate the length of stay for each encounter in the provided dictionary.
    If the primary diagnosis is "Observation", return 0 days for that encounter.

    >>> encounters = {
    ...     'E1234': {'admit date': date(2023, 1, 1), 'primary diagnosis': 'Flu', 'discharge date': date(2023, 1, 5)},
    ...     'E1235': {'admit date': date(2023, 1, 2), 'primary diagnosis': 'Observation', 'discharge date': date(2023, 1, 3)},
    ...     'E1236': {'admit date': date(2023, 1, 10), 'primary diagnosis': 'Covid', 'discharge date': date(2023, 1, 15)},
    ... }
    >>> length_of_stay(encounters)
    {'E1234': 4, 'E1235': 0, 'E1236': 5}
    """

    los = {}

    for encounter_id, data in encounters.items():
        # Ensure that necessary fields are present
        if all(key in data for key in ['admit date', 'discharge date', 'primary diagnosis']):
            admit_date = data['admit date']
            discharge_date = data['discharge date']
            diagnosis = data['primary diagnosis']

            if diagnosis == "Observation":
                los[encounter_id] = 0
            else:
                # Calculate length of stay
                stay_length = (discharge_date - admit_date).days
                los[encounter_id] = stay_length
        else:
            los[encounter_id] = "Missing required data"

    return los

```

```python
from datetime import date

def length_of_stay(encounters):
    """(dict) -> dict
    Calculate the length of stay for each encounter in the provided dictionary.
    If the primary diagnosis is "Observation", return 0 days for that encounter.

    >>> encounters = {
    ...     'E1234': {'admit date': date(2023, 1, 1), 'primary diagnosis': 'Flu', 'discharge date': date(2023, 1, 5)},
    ...     'E1235': {'admit date': date(2023, 1, 2), 'primary diagnosis': 'Observation', 'discharge date': date(2023, 1, 3)},
    ...     'E1236': {'admit date': date(2023, 1, 10), 'primary diagnosis': 'Covid', 'discharge date': date(2023, 1, 15)},
    ... }
    >>> length_of_stay(encounters)
    {'E1234': 4, 'E1235': 0, 'E1236': 5}
    """

    los = {}

    for encounter_id, data in encounters.items():
        # Ensure that necessary fields are present
        if all(key in data for key in ['admit date', 'discharge date', 'primary diagnosis']):
            admit_date = data['admit date']
            discharge_date = data['discharge date']
            diagnosis = data['primary diagnosis']

            if diagnosis == "Observation":
                los[encounter_id] = 0
            else:
                # Calculate length of stay
                stay_length = (discharge_date - admit_date).days
                los[encounter_id] = stay_length
        else:
            los[encounter_id] = "Missing required data"

    return los

# Example usage
if __name__ == "__main__":
    encounters = {
        'E1234': {'admit date': date(2023, 1, 1), 'primary diagnosis': 'Flu', 'discharge date': date(2023, 1, 5)},
        'E1235': {'admit date': date(2023, 1, 2), 'primary diagnosis': 'Observation', 'discharge date': date(2023, 1, 3)},
        'E1236': {'admit date': date(2023, 1, 10), 'primary diagnosis': 'Covid', 'discharge date': date(2023, 1, 15)},
    }

    print(length_of_stay(encounters))

```

Output:
```
{'E1234': 4, 'E1235': 0, 'E1236': 5}
```

```python
# Example usage
if __name__ == "__main__":
    encounters = {
        'E1234': {'admit date': date(2023, 1, 1), 'primary diagnosis': 'Flu', 'discharge date': date(2023, 1, 5)},
        'E1235': {'admit date': date(2023, 1, 2), 'primary diagnosis': 'Observation', 'discharge date': date(2023, 1, 3)},
        'E1236': {'admit date': date(2023, 1, 10), 'primary diagnosis': 'Covid', 'discharge date': date(2023, 1, 15)},
    }

    print(length_of_stay(encounters))
```

Output:
```
{'E1234': 4, 'E1235': 0, 'E1236': 5}
```

```python
import doctest
doctest.run_docstring_examples(length_of_stay, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    encounters = {
        'E1234': {'admit date': date(2023, 1, 1), 'primary diagnosis': 'Flu', 'discharge date': date(2023, 1, 5)},
        'E1235': {'admit date': date(2023, 1, 2), 'primary diagnosis': 'Observation', 'discharge date': date(2023, 1, 3)},
        'E1236': {'admit date': date(2023, 1, 10), 'primary diagnosis': 'Covid', 'discharge date': date(2023, 1, 15)},
    }
Expecting nothing
ok
Trying:
    length_of_stay(encounters)
Expecting:
    {'E1234': 4, 'E1235': 0, 'E1236': 5}
ok
```

```python
def length_of_stay(admit_date, discharge_date, diagnosis):
    """(date, date, str) -> int
    Computes the length of stay in days.
    If the diagnosis is "Observation", return 0.
    """
    if diagnosis == "Observation":
        return 0
    return (discharge_date - admit_date).days
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}
# Loop through encounters and call the function with the correct arguments
for e, v in encounters.items():
    print(f"Encounter {e}: Length of Stay = {length_of_stay(v['admit'], v['discharge'], v['diagnosis'])} days")
```

Output:
```
Encounter E1234: Length of Stay = 5 days
Encounter E8342: Length of Stay = 4 days
Encounter E9231: Length of Stay = 1 days
Encounter E8333: Length of Stay = 0 days
Encounter E3342: Length of Stay = 0 days
```

```python
### MY OWN TEST CASES 1.
from datetime import date

# Custom test cases
def custom_length_of_stay_tests():
    test_cases = {
        "Test Case 1": {  # Regular encounter
            "admit": date(2021, 6, 1),
            "discharge": date(2021, 6, 5),
            "diagnosis": "Pneumonia"
        },
        "Test Case 2": {  # Encounter with observation diagnosis
            "admit": date(2021, 6, 2),
            "discharge": date(2021, 6, 3),
            "diagnosis": "Observation"
        },
        "Test Case 3": {  # Same day admit and discharge
            "admit": date(2021, 6, 10),
            "discharge": date(2021, 6, 10),
            "diagnosis": "Flu"
        },
        "Test Case 4": {  # Longer stay
            "admit": date(2021, 5, 20),
            "discharge": date(2021, 6, 5),
            "diagnosis": "Surgery"
        },
        "Test Case 5": {  # Short stay with zero LOS
            "admit": date(2021, 7, 1),
            "discharge": date(2021, 7, 1),
            "diagnosis": "Check-up"
        }
    }

    for name, data in test_cases.items():
        los = length_of_stay(data['admit'], data['discharge'], data['diagnosis'])
        print(f"{name}: Length of Stay = {los} days")

# Run custom test cases
custom_length_of_stay_tests()

```

Output:
```
Test Case 1: Length of Stay = 4 days
Test Case 2: Length of Stay = 0 days
Test Case 3: Length of Stay = 0 days
Test Case 4: Length of Stay = 16 days
Test Case 5: Length of Stay = 0 days
```

```python
### MY OWN TEST CASES 2.
from datetime import date

def additional_length_of_stay_tests():
    test_cases = {
        "Test Case 1": {  # Regular encounter with a weekend
            "admit": date(2022, 2, 25),  # Friday
            "discharge": date(2022, 2, 27),  # Sunday
            "diagnosis": "Bronchitis"
        },
        "Test Case 2": {  # Extended stay for recovery
            "admit": date(2022, 1, 15),
            "discharge": date(2022, 2, 1),
            "diagnosis": "Recovery"
        },
        "Test Case 3": {  # No stay due to observation
            "admit": date(2022, 3, 10),
            "discharge": date(2022, 3, 10),
            "diagnosis": "Observation"
        },
        "Test Case 4": {  # Long-term stay
            "admit": date(2022, 5, 1),
            "discharge": date(2022, 6, 15),
            "diagnosis": "Heart Surgery"
        },
        "Test Case 5": {  # Stay with a holiday
            "admit": date(2022, 12, 24),  # Christmas Eve
            "discharge": date(2022, 12, 26),  # Christmas Day
            "diagnosis": "Holiday Check"
        }
    }

    for name, data in test_cases.items():
        los = length_of_stay(data['admit'], data['discharge'], data['diagnosis'])
        print(f"{name}: Length of Stay = {los} days")

# Run additional test cases
additional_length_of_stay_tests()

```

Output:
```
Test Case 1: Length of Stay = 2 days
Test Case 2: Length of Stay = 17 days
Test Case 3: Length of Stay = 0 days
Test Case 4: Length of Stay = 45 days
Test Case 5: Length of Stay = 2 days
```

---
---

### 5.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
test1 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

test2 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

### BEGIN SOLUTION

"""
>>> average_los(test1)
5
"""


### END SOLUTION
```

Output:
```
'\n>>> average_los(test1)\n5\n'
```

```python
from datetime import date

def length_of_stay(admit_date, discharge_date):
    """(date, date) -> int
    Computes the length of stay in days.
    """
    return (discharge_date - admit_date).days

def average_los(encounters):
    """(dict) -> float
    Computes the average length of stay (LOS) for the encounters.
    If the LOS for any encounter is 0, it is not counted in the average.

    >>> test1 = {
    ...     "E1234": {"admit": date(2019, 1, 3), "diagnosis": "COPD", "discharge": date(2019, 1, 8)},
    ... }
    >>> average_los(test1)
    5.0

    >>> test2 = {
    ...     "E1234": {"admit": date(2019, 1, 3), "diagnosis": "COPD", "discharge": date(2019, 1, 3)},
    ...     "E1235": {"admit": date(2019, 1, 4), "diagnosis": "Hypertension", "discharge": date(2019, 1, 8)},
    ... }
    >>> average_los(test2)
    4.0
    """
    total_los = 0
    count = 0

    for encounter in encounters.values():
        admit_date = encounter["admit"]
        discharge_date = encounter["discharge"]

        los = length_of_stay(admit_date, discharge_date)

        if los > 0:  # Only count non-zero lengths of stay
            total_los += los
            count += 1

    return total_los / count if count > 0 else 0.0
```

```python
# Example usage and doctest
if __name__ == "__main__":
    import doctest
    doctest.run_docstring_examples(average_los, globals(), verbose=True)

    # Additional test case
    test3 = {
        "E1236": {"admit": date(2019, 1, 10), "diagnosis": "Covid", "discharge": date(2019, 1, 15)},
        "E1237": {"admit": date(2019, 1, 12), "diagnosis": "Observation", "discharge": date(2019, 1, 12)},
    }
    print(average_los(test3))  # Expected: 5.
```

Output:
```
Finding tests in NoName
Trying:
    test1 = {
        "E1234": {"admit": date(2019, 1, 3), "diagnosis": "COPD", "discharge": date(2019, 1, 8)},
    }
Expecting nothing
ok
Trying:
    average_los(test1)
Expecting:
    5.0
ok
Trying:
    test2 = {
        "E1234": {"admit": date(2019, 1, 3), "diagnosis": "COPD", "discharge": date(2019, 1, 3)},
        "E1235": {"admit": date(2019, 1, 4), "diagnosis": "Hypertension", "discharge": date(2019, 1, 8)},
    }
Expecting nothing
ok
Trying:
    average_los(test2)
Expecting:
    4.0
ok
5.0
```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(average_los(encounters)) == 2.75
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    test1 = {
        "E1234": {"admit": date(2019, 1, 3), "diagnosis": "COPD", "discharge": date(2019, 1, 8)},
    }
Expecting nothing
ok
Trying:
    average_los(test1)
Expecting:
    5.0
ok
Trying:
    test2 = {
        "E1234": {"admit": date(2019, 1, 3), "diagnosis": "COPD", "discharge": date(2019, 1, 3)},
        "E1235": {"admit": date(2019, 1, 4), "diagnosis": "Hypertension", "discharge": date(2019, 1, 8)},
    }
Expecting nothing
ok
Trying:
    average_los(test2)
Expecting:
    4.0
ok
```

```python
## MY OWN TEST CASES 1.
from datetime import date

# Define custom test cases
def custom_average_los_tests():
    test_cases = {
        "Test Case 1": {  # Regular case with valid LOS
            "E1234": {"admit": date(2019, 1, 3), "diagnosis": "COPD", "discharge": date(2019, 1, 8)},
            "E1235": {"admit": date(2019, 1, 4), "diagnosis": "Hypertension", "discharge": date(2019, 1, 9)},
        },
        "Test Case 2": {  # Case with zero LOS
            "E1236": {"admit": date(2019, 1, 10), "diagnosis": "Flu", "discharge": date(2019, 1, 10)},
            "E1237": {"admit": date(2019, 1, 11), "diagnosis": "Cold", "discharge": date(2019, 1, 13)},
        },
        "Test Case 3": {  # Mixed valid and zero LOS
            "E1238": {"admit": date(2019, 1, 1), "diagnosis": "Pneumonia", "discharge": date(2019, 1, 5)},
            "E1239": {"admit": date(2019, 1, 6), "diagnosis": "Observation", "discharge": date(2019, 1, 7)},
        },
        "Test Case 4": {  # No valid LOS
            "E1240": {"admit": date(2019, 1, 10), "diagnosis": "Observation", "discharge": date(2019, 1, 10)},
        },
    }

    for name, encounters in test_cases.items():
        result = average_los(encounters)
        print(f"{name}: Average LOS = {result:.1f}")

# Run custom test cases
custom_average_los_tests()

```

Output:
```
Test Case 1: Average LOS = 5.0
Test Case 2: Average LOS = 2.0
Test Case 3: Average LOS = 2.5
Test Case 4: Average LOS = 0.0
```

```python
### MY OWN TEST CASES 2.
from datetime import date

# Define another set of custom test cases
def additional_average_los_tests():
    additional_cases = {
        "Test Case 1": {  # All valid encounters
            "E1250": {"admit": date(2020, 2, 1), "diagnosis": "Flu", "discharge": date(2020, 2, 5)},
            "E1251": {"admit": date(2020, 2, 3), "diagnosis": "Covid", "discharge": date(2020, 2, 7)},
            "E1252": {"admit": date(2020, 2, 5), "diagnosis": "Asthma", "discharge": date(2020, 2, 8)},
        },
        "Test Case 2": {  # All encounters with zero LOS
            "E1253": {"admit": date(2020, 3, 1), "diagnosis": "Observation", "discharge": date(2020, 3, 1)},
            "E1254": {"admit": date(2020, 3, 2), "diagnosis": "Observation", "discharge": date(2020, 3, 2)},
        },
        "Test Case 3": {  # Mixed valid LOS and zero LOS
            "E1255": {"admit": date(2020, 4, 1), "diagnosis": "Diabetes", "discharge": date(2020, 4, 4)},
            "E1256": {"admit": date(2020, 4, 5), "diagnosis": "Observation", "discharge": date(2020, 4, 5)},
            "E1257": {"admit": date(2020, 4, 6), "diagnosis": "Heart Disease", "discharge": date(2020, 4, 10)},
        },
        "Test Case 4": {  # Average with one encounter
            "E1258": {"admit": date(2020, 5, 1), "diagnosis": "Migraine", "discharge": date(2020, 5, 3)},
        },
    }

    for name, encounters in additional_cases.items():
        result = average_los(encounters)
        print(f"{name}: Average LOS = {result:.1f}")

# Run additional test cases
additional_average_los_tests()

```

Output:
```
Test Case 1: Average LOS = 3.7
Test Case 2: Average LOS = 0.0
Test Case 3: Average LOS = 3.5
Test Case 4: Average LOS = 2.0
```

---

### 5.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION

### END SOLUTION
```

```python
def c_to_f(celsius):
    """(float) -> float
    Converts degrees Celsius to degrees Fahrenheit.

    The formula for conversion is:
    temp_f = (temp_c * 9/5) + 32

    >>> c_to_f(100)
    212.0
    >>> c_to_f(0)
    32.0
    >>> c_to_f(37)
    98.6
    """
    return (celsius * 9/5) + 32



```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    c_to_f(37)
Expecting:
    98.6
ok
```

```python
### MY OWN TEST CASES 1.
def custom_c_to_f_tests():
    test_cases = {
        "Test Case 1": 100,     # Boiling point of water
        "Test Case 2": 0,       # Freezing point of water
        "Test Case 3": 37,      # Average human body temperature
        "Test Case 4": -40,     # Where Celsius and Fahrenheit are equal
        "Test Case 5": 20,      # Room temperature
        "Test Case 6": 25,      # Typical warm day
    }

    for name, celsius in test_cases.items():
        fahrenheit = c_to_f(celsius)
        print(f"{name}: {celsius}C = {fahrenheit:.1f}F")

# Run custom test cases
custom_c_to_f_tests()

```

Output:
```
Test Case 1: 100C = 212.0F
Test Case 2: 0C = 32.0F
Test Case 3: 37C = 98.6F
Test Case 4: -40C = -40.0F
Test Case 5: 20C = 68.0F
Test Case 6: 25C = 77.0F
```

```python
### MY OWN TESR CASES 2.
def additional_c_to_f_tests():
    test_cases = {
        "Test Case 1": 15,      # A cool day
        "Test Case 2": -10,     # Below freezing
        "Test Case 3": 50,      # Warm day
        "Test Case 4": 90,      # Hot day
        "Test Case 5": 37.5,    # Slightly above average body temperature
        "Test Case 6": -273.15,  # Absolute zero
    }

    for name, celsius in test_cases.items():
        fahrenheit = c_to_f(celsius)
        print(f"{name}: {celsius}C = {fahrenheit:.1f}F")

# Run additional test cases
additional_c_to_f_tests()

```

Output:
```
Test Case 1: 15C = 59.0F
Test Case 2: -10C = 14.0F
Test Case 3: 50C = 122.0F
Test Case 4: 90C = 194.0F
Test Case 5: 37.5C = 99.5F
Test Case 6: -273.15C = -459.7F
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.