<a href="https://colab.research.google.com/github/sravanthi200719/HDS5210_InClass/blob/master/week06_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In the /data directory on the Jupyter server, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Put your solutioimport json

ALLERGIES_FILE = "allergies.json"

def allergy_count(json_file):
    with open(json_file, 'r') as file:
        data = json.load(file)
    return len(data['entry'])

# Call the function and print the result
result = allergy_count(ALLERGIES_FILE)
print(f"Number of entries: {result}")


```

Output:
```
Number of entries: 4
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

```python
# Assertions to verify the result
assert type(result) == int, "The result should be an integer"
assert result == 4, "The number of entries should be 4"
```

my own test case for this with my file .

```python
import csv

CSV_FILE = "est2018-pop-res.csv"

def state_count(csv_file):
    with open(csv_file, 'r') as file:
        csv_reader = csv.DictReader(file)
        return sum(1 for row in csv_reader if row['SUMLEV'] == '040')

# Call the function and print the result
result = state_count(CSV_FILE)
print(f"Number of states (including DC and Puerto Rico): {result}")

```

Output:
```
Number of states (including DC and Puerto Rico): 52
```

```python
# Assertions to verify the result
assert type(result) == int, "The result should be an integer"
assert result == 52, "The number of states (including DC and Puerto Rico) should be 52"
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Put yimport json

ALLERGIES_FILE = "allergies.json"

def patient_count(json_file):
    with open(json_file, 'r') as file:
        data = json.load(file)

    patients = set()
    for entry in data['entry']:
        patient_reference = entry['resource']['patient']['reference']
        patients.add(patient_reference)

    return len(patients)

# Call the function and print the result
result = patient_count(ALLERGIES_FILE)
print(f"Number of unique patients: {result}")

```

Output:
```
Number of unique patients: 2
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### my own test case for above file .

```python
import csv

CSV_FILE = "est2018-pop-res.csv"

def unique_region_count(csv_file):
    with open(csv_file, 'r') as file:
        csv_reader = csv.DictReader(file)
        regions = set()
        for row in csv_reader:
            if row['SUMLEV'] == '040':  # Only count state-level entries
                regions.add(row['REGION'])
    return len(regions)

# Call the function and print the result
result = unique_region_count(CSV_FILE)
print(f"Number of unique regions: {result}")

```

Output:
```
Number of unique regions: 5
```

```python
# Assertions to verify the result
assert type(result) == int, "The result should be an integer"
assert result == 5, "The number of unique regions should be 4"

print("All assertions passed!")
```

Output:
```
All assertions passed!
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Put your solution hereimport json

ALLERGIES_FILE = "test_allergies.json"

import json

ALLERGIES_FILE = "allergies.json"

def allergy_per_patient(json_file):
    with open(json_file, 'r') as file:
        data = json.load(file)

    allergies = {}
    for entry in data['entry']:
        patient_name = entry['resource']['patient']['display']
        allergies[patient_name] = allergies.get(patient_name, 0) + 1

    return allergies

# Call the function and print the result
result = allergy_per_patient(ALLERGIES_FILE)
print("Allergies per patient:")
for patient, count in result.items():
    print(f"{patient}: {count}")


```

Output:
```
Allergies per patient:
Jason Argonaut: 3
Paul Boal: 1
```

```python
# Assertions to verify the result
assert type(result) == dict, "The result should be a dictionary"
assert result == {'Paul Boal': 1, 'Jason Argonaut': 3}, "The allergy counts are incorrect"

print("All assertions passed!")
```

Output:
```
All assertions passed!
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### my own test case for my file .

```python
import csv

CSV_FILE = "est2018-pop-res.csv"

def states_per_region(csv_file):
    with open(csv_file, 'r') as file:
        csv_reader = csv.DictReader(file)
        region_counts = {}
        for row in csv_reader:
            if row['SUMLEV'] == '040':  # Only count state-level entries
                region = row['REGION']
                region_counts[region] = region_counts.get(region, 0) + 1
    return region_counts

# Call the function and print the result
result = states_per_region(CSV_FILE)
print("States per region:")
for region, count in result.items():
    print(f"Region {region}: {count} states")

```

Output:
```
States per region:
Region 3: 17 states
Region 4: 13 states
Region 1: 9 states
Region 2: 12 states
Region X: 1 states
```

```python
# Assertions to verify the result
assert type(result) == dict, "The result should be a dictionary"
expected_result = {'3': 17, '4': 13, '1': 9, '2': 12, 'X':1}
assert result == expected_result, f"The state counts are incorrect. Expected {expected_result}, but got {result}"

print("All assertions passed!")
```

Output:
```
All assertions passed!
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
import json

ALLERGIES_FILE = "allergies.json"

def allergy_list(json_file):
    with open(json_file, 'r') as file:
        data = json.load(file)

    result = []
    for entry in data['entry']:
        patient_name = entry['resource']['patient']['display']
        substance = entry['resource']['substance']['text']
        reaction = entry['resource']['reaction'][0]['manifestation'][0]['text']
        result.append([patient_name, substance, reaction])

    return sorted(result)

# Call the function and print the result
result = allergy_list(ALLERGIES_FILE)
print("Allergy list:")
for item in result:
    print(f"{item[0]}: {item[1]} - {item[2]}")

```

Output:
```
Allergy list:
Jason Argonaut: PENICILLIN G - Hives
Jason Argonaut: SHELLFISH-DERIVED PRODUCTS - Itching
Jason Argonaut: STRAWBERRY - Anaphylaxis
Paul Boal: PENICILLIN G - Bruising
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### my own test case for my own file .

```python
import csv

CSV_FILE = "est2018-pop-res.csv"

def state_info_list(csv_file):
    with open(csv_file, 'r') as file:
        csv_reader = csv.DictReader(file)
        result = []
        for row in csv_reader:
            if row['SUMLEV'] == '040' and row['STATE'] != '00' and row['STATE'] != '72':  # Exclude US total and Puerto Rico
                state_name = row['NAME']
                region = row['REGION']
                division = row['DIVISION']
                population = int(row['POPESTIMATE2018'])
                result.append([state_name, region, division, population])

    return sorted(result)

# Call the function and print the result
result = state_info_list(CSV_FILE)
print("State information list:")
for item in result[:5]:  # Print first 5 items
    print(f"{item[0]}: Region {item[1]}, Division {item[2]}, Population {item[3]:,}")

```

Output:
```
State information list:
Alabama: Region 3, Division 6, Population 4,887,871
Alaska: Region 4, Division 9, Population 737,438
Arizona: Region 4, Division 8, Population 7,171,646
Arkansas: Region 3, Division 7, Population 3,013,825
California: Region 4, Division 9, Population 39,557,045
```

```python
# Assertion to verify the result
assert len(result) == 51, f"Expected 51 states (including DC), but got {len(result)}"
assert result[0][0] == "Alabama", f"First state should be Alabama, but got {result[0][0]}"
assert result[-1][0] == "Wyoming", f"Last state should be Wyoming, but got {result[-1][0]}"

print("Assertions passed!")
```

Output:
```
Assertions passed!
```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
import json

ALLERGIES_FILE = "allergies.json"

def allergy_list(json_file):
    with open(json_file, 'r') as file:
        data = json.load(file)

    result = []
    for entry in data['entry']:
        patient_name = entry['resource']['patient']['display']
        substance = entry['resource']['substance']['text']
        reaction = entry['resource']['reaction'][0]['manifestation'][0]['text']
        result.append([patient_name, substance, reaction])

    return result

def allergy_reaction(json_file, patient, substance):
    allergies = allergy_list(json_file)
    for allergy in allergies:
        if allergy[0] == patient and allergy[1] == substance:
            return allergy[2]
    return None
```

```python
# Test the function and print the results
print(f"Jason Argonaut's reaction to PENICILLIN G: {allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')}")
print(f"Jason Argonaut's reaction to SHELLFISH-DERIVED PRODUCTS: {allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS')}")
print(f"Jason Argonaut's reaction to STRAWBERRY: {allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY')}")
print(f"Jason Argonaut's reaction to PENICILLIN: {allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN')}")
print(f"Paul Boal's reaction to PENICILLIN G: {allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G')}")
```

Output:
```
Jason Argonaut's reaction to PENICILLIN G: Hives
Jason Argonaut's reaction to SHELLFISH-DERIVED PRODUCTS: Itching
Jason Argonaut's reaction to STRAWBERRY: Anaphylaxis
Jason Argonaut's reaction to PENICILLIN: None
Paul Boal's reaction to PENICILLIN G: Bruising
```

```python
# Assertions to verify the results
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'

print("All assertions passed!")
```

Output:
```
All assertions passed!
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

### my own test case for my file .

```python
import csv

CSV_FILE = "est2018-pop-res.csv"

def state_info_list(csv_file):
    with open(csv_file, 'r') as file:
        csv_reader = csv.DictReader(file)
        result = []
        for row in csv_reader:
            if row['SUMLEV'] == '040' and row['STATE'] != '00' and row['STATE'] != '72':  # Exclude US total and Puerto Rico
                state_name = row['NAME']
                region = row['REGION']
                population = int(row['POPESTIMATE2018'])
                result.append([state_name, region, population])

    return result

def state_population(csv_file, region, state):
    states = state_info_list(csv_file)
    for state_info in states:
        if state_info[1] == region and state_info[0] == state:
            return state_info[2]
    return None


```

```python
# Test the function and print the results
print(f"Population of California in Region 4: {state_population(CSV_FILE, '4', 'California'):,}")
print(f"Population of New York in Region 1: {state_population(CSV_FILE, '1', 'New York'):,}")
print(f"Population of Florida in Region 3: {state_population(CSV_FILE, '3', 'Florida'):,}")
print(f"Population of Texas in Region 2: {state_population(CSV_FILE, '2', 'Texas')}")
print(f"Population of Alaska in Region 4: {state_population(CSV_FILE, '4', 'Alaska'):,}")
```

Output:
```
Population of California in Region 4: 39,557,045
Population of New York in Region 1: 19,542,209
Population of Florida in Region 3: 21,299,325
Population of Texas in Region 2: None
Population of Alaska in Region 4: 737,438
```

```python
# Assertions to verify the results
assert state_population(CSV_FILE, '4', 'California') == 39557045
assert state_population(CSV_FILE, '1', 'New York') == 19542209
assert state_population(CSV_FILE, '3', 'Florida') == 21299325
assert state_population(CSV_FILE, '2', 'Texas') == None
assert state_population(CSV_FILE, '4', 'Alaska') == 737438

print("All assertions passed!")
```

Output:
```
All assertions passed!
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

